<a href="https://colab.research.google.com/github/lorcan2440/Computing-Exercises/blob/main/Copy_of_06_Exercises.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
def area(P_1: tuple, P_2: tuple, P_3: tuple):
  if not len(P_1) == len(P_2) == len(P_3) == 2:
    raise ValueError('All co-ordinates must be 2-tuples.')
  return abs((P_1[0] * (P_2[1] - P_3[1]) + P_2[0] * (P_3[1] - P_1[1]) + P_3[0] * (P_1[1] - P_2[1]))/2)

# Enter arguments to function as three 2-tuples
print(area((0, 0), (1, 1), (-1, 2)))
  
```

Output:
```
1.5
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
import math

def area(points: list):
  # Check that the input is the correct form
  if not all([type(x) is tuple and len(x) == 2 for x in points]):
    raise ValueError('Incorrect input syntax - enter a list of 2-tuples')

  # Get the centroid
  centroid = (sum(points[i][0] for i in range(len(points))) / len(points),
              sum(points[j][1] for j in range(len(points))) / len(points))

  # Sort the points clockwise from the origin by finding their argument relative to the centroid in polar coordinates
  points_list = points
  points_list.sort(key = lambda p: math.atan2(p[0] - centroid[0], p[1] - centroid[1]))
  
  # Use area formula
  area = 0
  for i in range(len(points_list) - 1):
    area += points_list[i][0] * points_list[i+1][1] - points_list[i+1][0] * points_list[i][1]
  area = abs(area / 2)
  return area

# House shaped pentagon, un-ordered points to test sorting
print(area([(0, 0), (2, 0), (1, 2), (2, 1), (0, 1)]))

```

Output:
```
3.0
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x: list, y: list):
  # Return sum of two vectors
  if not len(x) == len(y):
    raise ValueError('Each input must be a list of equal length')
  vector = []
  for i in range(len(x)):
    vector.append(x[i] + y[i])
  return vector

```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Extension: list comprehension

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
def sum_vector(x, y):
  # Return sum of two vectors
  return [sum(list(zip(x, y))[i]) for i in range(len(x))]
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A dictionary from college abbreviation to name; and
1. A list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
colleges = {"Emmanuel" : "EM",
            "Christ's" : "CHR",
            "Downing" : "DOW",
            "Fitzwilliam" : "F",
            "Girton" : "G"}

sorted_colleges = {k: v for k, v in sorted(colleges.items(), key = lambda item: item[1])}
print(sorted_colleges)


```

Output:
```
{"Christ's": 'CHR', 'Downing': 'DOW', 'Emmanuel': 'EM', 'Fitzwilliam': 'F', 'Girton': 'G'}
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
colleges = {"Emmanuel" : ("EM", 1584, 718),
            "Christ's" : ("CHR", 1505, 689),
            "Downing" : ("DOW", 1800, 845),
            "Fitzwilliam" : ("F", 1869, 899),
            "Girton" : ("G", 1869, 808)}

by_date = {k: v for k, v in sorted(colleges.items(), key = lambda item: item[1][1])}
by_student_numbers = {k: v for k, v in sorted(colleges.items(), key = lambda item: item[1][2])}

most_students = list(by_student_numbers)[-1]
oldest_college = list(by_date)[0]
print('College with greatest number of students is {} ({})'.format(most_students, colleges[most_students][0]))
print('Oldest college is {} ({}), which has {} students'.format(oldest_college, colleges[oldest_college][0], colleges[oldest_college][2]))
```

Output:
```
College with greatest number of students is Fitzwilliam (F)
Oldest college is Christ's (CHR), which has 689 students
```