<a href="https://colab.research.google.com/github/lorcan2440/Computing-Exercises/blob/main/Copy_of_03_Exercises.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Exercise 03.1

Compare the computed values of 

$$ 
d_0 = a \cdot b + a  \cdot  c
$$

and 

$$ 
d_1 = a \cdot (b + c)
$$

when $a = 100$, $b = 0.1$ and $c = 0.2$. Store $d_{0}$ in the variable `d0` and $d_{1}$ in the variable `d1`.


Try checking for equality, e.g. `print(d0 == d1)`. 

<!-- grade_id: cell-447110c225124fc4, points: 0, grade: false, solution: true -->
```python
a, b, c = 100, 0.1, 0.2
d0 = a * b + a * c
d1 = a * (b + c)
print(d0)
print(d1)
print(d0 == d1)
```

Output:
```
30.0
30.000000000000004
False
```

<!-- grade_id: cell-91f2d89d88ea86c0, points: 0, grade: true, solution: false -->
```python
assert d0 == 30.0
assert d1 != 30.0
assert d0 != d1
```

## Exercise 03.2

For the polynomial 
\begin{align}
f(x, y) &= (x + y)^{6} 
\\
&=  x^6 + 6x^{5}y + 15x^{4}y^{2} + 20x^{3}y^{3} + 15x^{2}y^{4} + 6xy^{5} + y^{6}
\end{align}
compute $f$ using: (i) the compact form $(x + y)^{6}$; and (ii) the expanded form for:

(a) $x = 10$ and $y = 10.1$

(b) $x = 10$ and $y = -10.1$

and compare the number of significant digits for which the answers are the same.
Store the answer for the compact version using the variable `f0`, and using the variable `f1` for the expanded version.

For case (b), compare the computed and analytical solutions and consider the relative error.
Which approach would you recommend for computing this expression?

#### (a) $x = 10$ and $y = 10.1$

<!-- grade_id: cell-e8428e3de87d5500, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = 10.1

f0 = (x + y)**6
f1 = x**6 + 6*x**5 * y + 15*x**4 * y**2 + 20*x**3 * y**3 + 15 * x**2 * y**4 + 6*x*y**5 + y**6
i = 0
decimal_difference = False
try:
  while str(f0)[i] == str(f1)[i] or i > 100:
    if not str(f0)[i] == '.':
      decimal_difference = True
    i += 1

  if decimal_difference:
    i -= 1
except IndexError:
  pass

print(f0, f1)
print('Equivalent to {} significant digits'.format(i))
```

Output:
```
65944160.60120103 65944160.601201
Equivalent to 15 significant digits
```

<!-- grade_id: cell-55d184514edf3a67, points: 0, grade: true, solution: false -->
```python
import math
assert math.isclose(f0, 65944160.60120103, rel_tol=1e-10)
assert math.isclose(f1, 65944160.601201, rel_tol=1e-10)
```

#### (b) $x = 10$ and $y = -10.1$

<!-- grade_id: cell-1561d20281ff67f6, points: 0, grade: false, solution: true -->
```python
x = 10.0
y = -10.1

f0 = (x + y)**6
f1 = x**6 + 6*x**5 * y + 15*x**4 * y**2 + 20*x**3 * y**3 + 15 * x**2 * y**4 + 6*x*y**5 + y**6
i = 0
decimal_difference = False
try:
  while str(f0)[i] == str(f1)[i] or i > 100:
    if not str(f0)[i] == '.':
      decimal_difference = True
    i += 1

  if decimal_difference:
    i -= 1
except IndexError:
  pass

print(f0, f1)
print('Equivalent to {} significant digits'.format(i))
print('Error in compact form: {}'.format(abs(f0 - (-0.1)**6)))
print('Error in expanded form: {}'.format(abs(f1 - (-0.1)**6)))
print('Compact form is much more accurate')
```

Output:
```
9.999999999999788e-07 9.958166629076004e-07
Equivalent to 2 significant digits
Error in compact form: 2.159934015498466e-20
Error in expanded form: 4.1833370923999754e-09
Compact form is much more accurate
```

```python
import math
assert math.isclose(f0, 1.0e-6, rel_tol=1e-10)
assert math.isclose(f1, 1.0e-6, rel_tol=1e-2)
```

## Exercise 03.3

Consider the expression

$$
f = \frac{1}{\sqrt{x^2 - 1} - x}
$$

When $x$ is very large, the denominator approaches zero, which can cause problems.

Try rephrasing the problem and eliminating the fraction by multiplying the numerator and denominator by $\sqrt{x^2 - 1} + x$ and evaluate the two versions of the expression when:

(a) $x = 1 \times 10^{7}$

(b) $x = 1 \times 10^{9}$ (You may get a Python error for this case. Why?)

#### (a) $x = 1 \times 10^{7}$

<!-- grade_id: cell-149ac8791e7ff020, points: 0, grade: true, solution: true -->
```python
f = lambda x: 1/(math.sqrt(x**2 - 1) - x)
f_modified = lambda x: -(math.sqrt(x**2 - 1) + x)

print('Original function value: {}'.format(f(1e7)))
print('Manipulated function value: {}'.format(f_modified(1e7)))
```

Output:
```
Original function value: -19884107.85185185
Manipulated function value: -19999999.999999948
```

#### (b) $x = 1 \times 10^{9}$

<!-- grade_id: cell-4ef986ef1f5838ba, points: 0, grade: true, solution: true -->
```python
f = lambda x: 1/(math.sqrt(x**2 - 1) - x)
f_modified = lambda x: -(math.sqrt(x**2 - 1) + x)

try:
  print('Original function value: {}'.format(f(1e9)))
  print('Manipulated function value: {}'.format(f_modified(1e9)))
except ZeroDivisionError:
  print('There was an error since when x^2 is calculated, its value is 10^18.') 
  print('When 1 is subtracted, the resulting number has more than 16 significant digits so cannot be stored exactly using 32 bit arithmetic.')
  print('So then 10^18 and 10^18 - 1 are stored as the exact same number so their difference is zero, leading to division by zero error.')
  print('Only the modified expression can be used since it does not involve dividing by the "zero".')
  print('Manipulated function value: {}'.format(f_modified(1e7)))


```

Output:
```
There was an error since when x^2 + 1 is calculated, its value is 10^18.
When 1 is subtracted, the resulting number has more than 16 significant digits so cannot be stored exactly using 32 bit arithmetic.
So then 10^9 and 10^9 - 1 are stored as the exact same number so their difference is zero, leading to division by zero error.
Only the modified expression can be used since it does not involve dividing by the "zero".
Manipulated function value: -19999999.999999948
```