<a href="https://colab.research.google.com/github/lorcan2440/Computing-Exercises/blob/main/Copy_of_09_Exercises.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n: int): 
  "Compute the nth Fibonacci number using recursion"
  if n == 0:
    return 0
  elif n == 1:
    return 1
  elif n > 1:
    return f(n - 1) + f(n - 2)
  else:
    raise ValueError('Input must be an integer greater than or equal to zero')
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    "Evaluate polynomial function"
    return x**3 - 6*x**2 + 4*x + 12

def compute_root(my_f, x0, x1, tol, max_it):
  f = lambda x: my_f(x)
  i = 0
  a = x0
  b = x1
  if a > b:
    a, b = b, a
  f_mid = 1 + tol
  while i < max_it and abs(f_mid) > tol:
    x_mid = (a + b)/2
    f_0 = f(a)
    f_mid = f(x_mid)

    if f_0 * f_mid < 0:
      b = x_mid
    else:
      a = x_mid
    i += 1
  if i == max_it and abs(f_mid) > tol:
    raise RuntimeError('Bisection algorithm did not converge fast enough.')

  return x_mid, f_mid, i
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=20)
```